# Statystyka matematyczna - ćwiczenia laboratoryjne

Ten notatnik zalicza się do grupy zestawów zadań, na podstawie których odbywa się zaliczenie ćwiczeń i podlega zwrotowi do oceny w ustalonym na zajęciach terminie.

Uwagi i wytyczne ogólne dotyczące uzupełniania notatnika:
- Podczas wykonywania zadań należy korzystać wyłącznie z pakietów zaimportowanych na początku notatnika oraz z pakietów wchodzących w skład standardowej biblioteki Pythona, które można zaimportować samodzielnie we wskazanej komórce notatnika.
- Swoje rozwiązania należy wprowadzać wyłącznie w miejce następujących fragmentów kodu:<br/> `# YOUR CODE HERE`<br/> `raise NotImplementedError()`<br/> 
a odpowiedzi tekstowe w komórkach oznaczonych hasłem:<br/> 
`YOUR ANSWER HERE`<br/> 
Nie należy w żaden sposób modyfikować pozostałych fragmentów kodu oraz innych elementów notatnika, w szczególności dodawać lub usuwać komórek oraz zmieniać nazwy pliku.
- Otrzymywane wyniki i odpowiedzi mają być rezultatem wykonania napisanego kodu.
- Zadanie należy wykonać w taki sposób, aby podczas wykonywania kodu nie zostały wyświetlone żadne ostrzeżenia.
- Zawarte w notatniku automatyczne testy mają charakter poglądowy. Dotyczą one wybranych aspektów zadań i mają za zadanie wyłapać podstawowe błędy. Przejście wszystkich testów nie oznacza, że zadanie jest wykonane w całości poprawnie.

Uwagi i wytyczne ogólne dotyczące wizualizacji wyników:
- Wszystkie wykresy powinny być wykonane w jednolitym, przejrzystym i czytelnym stylu, posiadać odpowiednio dobrane proporcje i zakresy wartości osi.
- Wykresy oraz ich osie powinny mieć nadane tytuły. Jeżeli w obrębie figury znajduje się więcej niż jeden wykres to figura również powinna mieć nadany tytuł. 
- Figury powinny mieć ustawione białe tło, tak, aby niezależnie od ustawień notatnika wszystkie elementy wykresów były dobrze widoczne (domyślnie tło jest przeźroczyste co może powodować problemy w notatnikach z ustawionym ciemnym tłem).
- Rozmiar poziomy figur nie powinien przekraczać 20 cali.

Przed odesłaniem zestawu zadań do oceny proszę uzupełnić komórkę z danymi autorów rozwiązania (nazwa zespołu oraz imiona, nazwiska i numery indeksów członków zespołu) oraz upewnić się, że notatnik działa zgodnie z oczekiwaniami. W tym celu należy skorzystać z opcji **Restart Kernel and Run All Cells...** dostępnej na górnej belce notatnika pod symbolem $\blacktriangleright\blacktriangleright$.

Nazwa zespołu:
Członkowie:

---

# Zestaw zadań 2: Rozkłady zmiennych losowych ciągłych

```python
import numpy as np
import pandas as pd
import scipy as sp
import matplotlib as mpl
import matplotlib.pyplot as plt
import seaborn as sns
```

```python
# Miejsce do importu pakietów wchodzących w skład standardowej biblioteki Pythona oraz ustawienie opcji wykorzystywanych pakietów
```

### Funkcje do wykorzystania

```python
def vizualize_continuous_distribution(pdf, cdf, interval=None):
    """
    Parameters
    -------
    pdf: pandas.DataFrame
        Tablica zawierająca index z wariantami zmiennej losowej i nadaną nazwą zmiennej (np. x) oraz pojedynczą kolumnę
        z wartościami gęstościami prawdopodobieństwa wyliczonymi dla poszczególnych wariantów o nadanej nazwie bazującej na nazwie zmiennej (np. f(x)).
    cdf: pandas.DataFrame
        Tablica zawierająca index z wariantami zmiennej losowej i nadaną nazwą zmiennej (np. x) oraz pojedynczą kolumnę
        z wartościami dystrubuanty wyliczonymi poszczególnych wariantów o nadanej nazwie bazującej na nazwie zmiennej (np. F(x)).
    interval: list
        Lista zawierająca trzy elementy w następującej kolejności:
         - wartość dolnej granicy przedziału,
         - wartość górnej granicy przedziału,
         - "inner" lub "outer" w zależności od tego, czy ma być zakolorowana wewnętrzna, czy zewnętrzna część przedziału.
    """
    pdf_name= pdf.columns[0]
    cdf_name = cdf.columns[0]    
    
    fig, axes = plt.subplots(1, 2, figsize=(10, 5), facecolor='white')

    fig.suptitle("Rozkład prawdopodobieństwa i dystrybuanta zmiennej losowej")

    axes[0].plot(pdf.index, pdf[pdf_name], color="grey")
    axes[0].set_title("Gęstość prawdopodobieństwa")
    axes[0].set_xlabel(pdf.index.name)
    axes[0].set_ylabel(pdf_name)
    
    if interval is not None:
        if interval[2]=="inner":
            axes[0].fill_between(pdf.index, pdf[pdf_name], 0, where=((pdf.index > interval[0]) & (pdf.index<interval[1])), color='green', alpha=0.5)
        elif interval[2]=="outer":
            axes[0].fill_between(pdf.index, pdf[pdf_name], 0, where=((pdf.index < interval[0]) | (pdf.index>interval[1])), color='green', alpha=0.5)
    
    axes[1].plot(cdf.index, cdf[cdf_name], color="grey")
    axes[1].set_title("Dystrybuanta")
    axes[1].set_xlabel(cdf.index.name)
    axes[1].set_ylabel(cdf_name)
```

### Zadanie 1: Przykładowy ciągły rozkład prawdopodobieństwa  [10 pkt]

Zmienna losowa X ma funkcję gęstości prawdopodobieństwa daną wzorem:

$$f(x) = \begin{cases}
    0\text{ dla }x<-1\\
    0.75 \cdot (1-x^2)\text{ dla }|x|<=1\\
    0\text{ dla }x>1
    \end{cases}$$

#### a) Tablica wybranych wartości funkcji gęstości prawdopodobieństwa

Przygotuj dane niezbędne do przedstawienia funkcji gęstości prawdopodobieństwa zmiennej losowej X w formie tablicy `r1p` o następującej strukturze:

> Indeks:</br>
> x - 101 wartości zmiennej losowej równo rozłożonych pomiędzy wartościami -2 a 2, tak aby dwie wartości wypadały na granicach określonych we wzorze przedziałów. </br>

> Kolumny: </br>
> f(x) - gęstość prawdopodobieństwa wyliczona dla poszczególnych wartości x.

Uwagi do wykonania zadania:
 - Warianty przyjmowane przez zmienną z wygeneruj za pomocą funkcji `numpy.linspace()` lub `numpy.arange()`.

<!-- grade_id: cell-a9ebbffdc209fdc5, points: 0, grade: false, solution: true -->
```python
r_variants = np.linspace(-2,2,101)
r_pdf = np.zeros_like(r_variants)
mask_1 = np.all([r_variants>=-1, r_variants<=1], axis=0)
r_pdf[mask_1] = 0.75*(1-(r_variants[mask_1]*r_variants[mask_1]))
r1p = pd.DataFrame(data=r_pdf, index=r_variants,columns=["f(x)"])
r1p.index.name = 'x'
r1p
```

Output:
```
f(x)
x          
-2.00   0.0
-1.96   0.0
-1.92   0.0
-1.88   0.0
-1.84   0.0
...     ...
 1.84   0.0
 1.88   0.0
 1.92   0.0
 1.96   0.0
 2.00   0.0

[101 rows x 1 columns]
```

<!-- grade_id: cell-faf98a9d8c9da9db, points: 2, grade: true, solution: false -->
```python
# Komórka testowa
assert type(r1p) == pd.DataFrame
assert r1p.shape == (101, 1)
assert list(r1p.columns) == ["f(x)"]
assert r1p.index.name == 'x'
assert np.isclose(r1p.index.min(), -2)
assert np.isclose(r1p.index.max(), 2)
assert np.isclose(r1p["f(x)"].min(), 0.0)
assert np.isclose(r1p["f(x)"].max(), 0.75)
assert np.isclose(r1p["f(x)"].mean(), 0.24742574257425745)
```

#### b) Tablica wybranych wartości dystrybuanty

Przygotuj dane niezbędne do przedstawienia dystrybuanty zmiennej losowej X w formie tablicy `r1f` o następującej strukturze:

> Indeks:</br>
> x - 101 wartości zmiennej losowej równo rozłożonych pomiędzy wartościami -2 a 2,  tak aby dwie wartości wypadały na granicach określonych we wzorze przedziałów. </br>

> Kolumny: </br>
> F(x) - wartości dystrybuanty wyliczone dla poszczególnych wartości x.

Uwagi do wykonania zadania:
 - Warianty przyjmowane przez zmienną z wygeneruj za pomocą funkcji `numpy.linspace()` lub `numpy.arange()`.
 - Wartości dystrybuanty oblicz za pomocą funkcji `scipy.integrate.quad()`.

<!-- grade_id: cell-6d6e536b2df352d5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f5ef96e0af718a2e, points: 2, grade: true, solution: false -->
```python
# Komórka testowa
assert type(r1f) == pd.DataFrame
assert r1f.shape == (101, 1)
assert list(r1f.columns) == ["F(x)"]
assert r1f.index.name == 'x'
assert np.isclose(r1f.index.min(), -2)
assert np.isclose(r1f.index.max(), 2)
assert np.isclose(r1f["F(x)"].min(), 0.0)
assert np.isclose(r1f["F(x)"].max(), 1.0)
assert np.isclose(r1f["F(x)"].mean(), 0.5)
```

#### c) Wizualizacja funkcji gęstości prawdopodobieństwa i dystrybuanty
Wykorzystaj funkcję `vizualize_continuous_distribution()` do przedstawienia rozkładu prawdopodobieństwa oraz dystrybuanty rozkładu zmiennej losowej X w formie wykresów.

<!-- grade_id: cell-be8ee275fecd4171, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### d) Podstawowe parametry rozkładu

Oblicz i zapisz do zmiennych następujące parametry rozkładu zmiennej losowej X - wartość średnia (`r1_mean`), wariancja(`r1_var`) i odchylenie standardowe (`r1_std`).

<!-- grade_id: cell-92cd90185bb1294e, points: 0, grade: false, solution: true -->
```python
# Wartość oczekiwana
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4f26c0a8fc7528c0, points: 0, grade: false, solution: true -->
```python
# Wariancja
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-bcb8dd2d877e7470, points: 0, grade: false, solution: true -->
```python
# Odchylenie standardowe
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d73934266ccfc415, points: 5, grade: true, solution: false -->
```python
# Komórka testowa
assert np.isclose(r1_mean, 0.0)
assert np.isclose(r1_var, 0.2)
assert np.isclose(r1_std, 0.4472135954999579)
```

### Zadanie 2: Rozkład normalny [10 pkt]

Zmienna losowa Y ma rozkład normalny z wartością oczekiwaną $\mu=5$ i odchyleniem standardowym $\sigma=3.5$.

#### a) Tablica wybranych wartości funkcji gęstości prawdopodobieństwa

Przygotuj dane niezbędne do przedstawienia funkcji gęstości prawdopodobieństwa zmiennej losowej Y w formie tablicy `r2p` o następującej strukturze:

> Indeks:</br>
> y - 100 wartości zmiennej losowej równo rozłożonych pomiędzy wartościami, dla których dystrybuanta zmiennej losowej Y przyjmuje wartości 0.0001 i 0.9999</br>

> Kolumny: </br>
> f(y) - gęstość prawdopodobieństwa wyliczona dla poszczególnych wartości y.


Uwagi do wykonania zadania:
 - Warianty przyjmowane przez zmienną z wygeneruj za pomocą funkcji `numpy.linspace()` lub `numpy.arange()`.
 - Wartości gęstości prawdopodobieństwa oblicz za pomocą odpowiedniej funkcji z modułu `scipy.stats.norm`.

<!-- grade_id: cell-fef0353657c6ff67, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-aae34e891f30a871, points: 3, grade: true, solution: false -->
```python
# Komórka testowa
assert type(r2p) == pd.DataFrame
assert r2p.shape == (100, 1)
assert list(r2p.columns) == ["f(y)"]
assert r2p.index.name == 'y'
assert np.isclose(r2p.index.min(), -8.016557699094882)
assert np.isclose(r2p.index.max(), 18.01655769909498)
assert np.isclose(r2p["f(y)"].min(), 0.00011309941907425523)
assert np.isclose(r2p["f(y)"].max(), 0.11390311079919879)
assert np.isclose(r2p["f(y)"].mean(), 0.03802196004604897)
```

#### b) Tablica wybranych wartości dystrybuanty

Przygotuj dane niezbędne do przedstawienia funkcji gęstości prawdopodobieństwa oraz dystrybuanty zmiennej losowej Y w formie tablicy `r2f` o następującej strukturze:

> Indeks:</br>
> y - 100 wartości zmiennej losowej równo rozłożonych pomiędzy wartościami Y, dla których dystrybuanta zmiennej losowej Y przyjmuje wartości 0.0001 i 0.9999</br>

> Kolumny: </br>
> F(y) - wartości dystrybuanty wyliczone dla poszczególnych wartości y.


Uwagi do wykonania zadania:
 - Warianty przyjmowane przez zmienną Y wygeneruj za pomocą funkcji `numpy.linspace()` lub `numpy.arange()`.
 - Wartości dystrybuanty oblicz za pomocą odpowiedniej funkcji z modułu `scipy.stats.norm`.

<!-- grade_id: cell-67e4191c6c4dca0f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2ea0b0a81b04f2bd, points: 3, grade: true, solution: false -->
```python
# Komórka testowa
assert type(r2f) == pd.DataFrame
assert r2f.shape == (100, 1)
assert list(r2f.columns) == ["F(y)"]
assert r2f.index.name == 'y'
assert np.isclose(r2f.index.min(), -8.016557699094882)
assert np.isclose(r2f.index.max(), 18.01655769909498)
assert np.isclose(r2f["F(y)"].min(), 0.0001)
assert np.isclose(r2f["F(y)"].max(), 0.9999)
assert np.isclose(r2f["F(y)"].mean(), 0.5)
```

#### c) Wizualizacja funkcji gęstości prawdopodobieństwa i dystrybuanty
Wykorzystaj funkcję `vizualize_continuous_distribution()` do przedstawienia rozkładu prawdopodobieństwa oraz dystrybuanty rozkładu zmiennej losowej Y w formie wykresów.

<!-- grade_id: cell-3dc22544508bd329, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### d) Podstawowe parametry rozkładu

Oblicz i zapisz do zmiennych następujące parametry rozkładu zmiennej losowej Y - wartość średnia (`r2_mean`), wariancja(`r2_var`) i odchylenie standardowe (`r2_std`).

Uwagi do wykonania zadania:
 - Wartości parametrów oblicz za pomocą odpowiedniej funkcji z modułu `scipy.stats.norm`.

<!-- grade_id: cell-e27ded4e2ef48f81, points: 0, grade: false, solution: true -->
```python
# Wartość oczekiwana
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a3ddf2ff0ea42a6a, points: 0, grade: false, solution: true -->
```python
# Wariancja
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c9027c5caa419b87, points: 0, grade: false, solution: true -->
```python
# Odchylenie standardowe
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ced38983179c1067, points: 3, grade: true, solution: false -->
```python
# Komórka testowa
assert np.isclose(r2_mean, 5.0)
assert np.isclose(r2_var, 12.25)
assert np.isclose(r2_std, 3.5)
```

### Zadanie 3: Rozkład gamma [10 pkt]
Zmienna losowa Z ma rozkład gamma z parametram $k=4$ i parametrem $\theta=2$.

#### a) Tablica wybranych wartości funkcji gęstości prawdopodobieństwa

Przygotuj dane niezbędne do przedstawienia funkcji gęstości prawdopodobieństwa zmiennej losowej Z w formie tablicy `r3p` o następującej strukturze:

> Indeks:</br>
> z - 100 wartości zmiennej losowej równo rozłożonych od 0 do wartości z, dla której dystrybuanta zmiennej losowej Z przyjmuje wartość 0.9999.</br>

> Kolumny: </br>
> f(z) - gęstość prawdopodobieństwa wyliczona dla poszczególnych wartości z. </br>

Uwagi do wykonania zadania:
 - Warianty przyjmowane przez zmienną z wygeneruj za pomocą funkcji `numpy.linspace()` lub `numpy.arange()`.
 - Wartości gęstości prawdopodobieństwa oblicz za pomocą odpowiedniej funkcji z modułu `scipy.stats.gamma`.

<!-- grade_id: cell-569cd0a4a96cc523, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-345298e7155dbe64, points: 3, grade: true, solution: false -->
```python
# Komórka testowa
assert type(r3p) == pd.DataFrame
assert r3p.shape == (100, 1)
assert list(r3p.columns) == ["f(z)"]
assert r3p.index.name == 'z'
assert np.isclose(r3p.index.min(), 0.0)
assert np.isclose(r3p.index.max(), 31.827628001262585)
assert np.isclose(r3p["f(z)"].min(), 0.0)
assert np.isclose(r3p["f(z)"].max(), 0.11196678941579825)
assert np.isclose(r3p["f(z)"].mean(), 0.031102170912604324)
```

#### b) Tablica wybranych wartości dystrybuanty

Przygotuj dane niezbędne do przedstawienia dystrybuanty zmiennej losowej Z w formie tablicy `r3f` o następującej strukturze:

> Indeks:</br>
> z - 100 wartości zmiennej losowej równo rozłożonych od 0 do wartości z, dla której dystrybuanta zmiennej losowej Z przyjmuje wartość 0.9999.</br>

> Kolumny: </br>
> F(z) - wartości dystrybuanty wyliczone dla poszczególnych wartości z.

Uwagi do wykonania zadania:
 - Warianty przyjmowane przez zmienną Y wygeneruj za pomocą funkcji `numpy.linspace()` lub `numpy.arange()`.
 - Wartości dystrybuanty oblicz za pomocą odpowiedniej funkcji z modułu `scipy.stats.gamma`.

<!-- grade_id: cell-564615e4223c86a5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0ebfec21daf353ce, points: 3, grade: true, solution: false -->
```python
# Komórka testowa
assert type(r3f) == pd.DataFrame
assert r3f.shape == (100, 1)
assert list(r3f.columns) == ["F(z)"]
assert r3f.index.name == 'z'
assert np.isclose(r3f.index.min(), 0.0)
assert np.isclose(r3f.index.max(), 31.827628001262585)
assert np.isclose(r3f["F(z)"].min(), 0.0)
assert np.isclose(r3f["F(z)"].max(), 0.9999)
assert np.isclose(r3f["F(z)"].mean(), 0.7461665542590229)
```

#### c) Wizualizacja funkcji gęstości prawdopodobieństwa i dystrybuanty
Wykorzystaj funkcję `vizualize_continuous_distribution()` do przedstawienia rozkładu prawdopodobieństwa oraz dystrybuanty rozkładu zmiennej losowej Z w formie wykresów.

<!-- grade_id: cell-dd806296fa930e42, points: 1, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

#### d) Podstawowe parametry rozkładu

Oblicz i zapisz do zmiennych następujące parametry rozkładu zmiennej losowej Z - wartość średnia (`r3_mean`), wariancja(`r3_var`) i odchylenie standardowe (`r3_std`).

Uwagi do wykonania zadania:
 - Wartości parametrów oblicz za pomocą odpowiedniej funkcji z modułu `scipy.stats.gamma`.

<!-- grade_id: cell-6f7b223a4f1ec575, points: 0, grade: false, solution: true -->
```python
# Wartość oczekiwana
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-0f78954428f22a7d, points: 0, grade: false, solution: true -->
```python
# Wariancja
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-e6f70c1fdba8fbb6, points: 0, grade: false, solution: true -->
```python
# Odchylenie standardowe
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-a25664440a360f64, points: 3, grade: true, solution: false -->
```python
# Komórka testowa
assert np.isclose(r3_mean, 8.0)
assert np.isclose(r3_var, 16.0)
assert np.isclose(r3_std, 4.0)
```

### Zadanie 4: Wyliczanie i wizualizacja prawdopodobieństw [6 pkt]

Oblicz i przedstaw na wykresach następujące prawdopodobieństwa:
- P(X>-0.5)
- P(1<Y<8)
- P(Z<2 | Z>12)

Wyniki zapisz kolejno do zmiennych `p1`, `p2` i `p3` oraz przedstaw na wykresach korzystając z funkcji `vizualize_continuous_distribution()`.

#### a) P(X>-0.5)

<!-- grade_id: cell-d83ae2882074819a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-d15d12889de658af, points: 2, grade: true, solution: false -->
```python
# Komórka testowa
assert np.isclose(p1, 0.84375)
```

#### b) P(1<Y<8)

<!-- grade_id: cell-241661cf0a3e5369, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-90a814e79dcf7ad1, points: 2, grade: true, solution: false -->
```python
# Komórka testowa
assert np.isclose(p2, 0.6777680763726662)
```

#### c) P(Z<2 | Z>12)

<!-- grade_id: cell-9069cc83e52cffd0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1e1be63339d333a8, points: 2, grade: true, solution: false -->
```python
# Komórka testowa
assert np.isclose(p3, 0.17019203965280166)
```